<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="coderunner">
    <name><text>PROTOTYPE_go_program</text></name>
    <questiontext format="html"><text><![CDATA[<p>Prototype for Go WAT questions. The learner submits a complete Go source file containing <code>package main</code>, any imports, and their functions, but no <code>main</code> function. Each test case's testcode is Go statements compiled into a generated harness main, selected by argument, one execution per test. Test code may use <code>fmt</code> and the learner's identifiers only.</p>
<p>Template parameters (JSON), all optional and space-separated: <code>banned</code> tokens that must not appear, <code>required</code> tokens that must appear, <code>limit</code> countable bounds (<code>return=N</code>, <code>bodylines=N</code>). Tokens are keywords or identifiers (word-boundary match on comment- and literal-stripped source), selectors like <code>errors.Is</code>, operators from the fixed table (<code>== != + &amp; &amp;&amp; || ! &lt;- := % *</code>), the pseudo-token <code>nakedreturn</code>, or literals (digit sequences, a quote character, or a <code>%</code> verb such as <code>%w</code>) matched on comment-stripped source. A test case whose testcode is exactly <code>// rules</code> prints <code>rules ok</code> when every check passes. On a violation the wrapper prints the reason as the first test's output and skips compilation.</p>]]></text></questiontext>
    <generalfeedback format="html"><text></text></generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <coderunnertype>go_program</coderunnertype>
    <prototypetype>1</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime>0</penaltyregime>
    <precheck>0</precheck>
    <showsource>0</showsource>
    <answerboxlines>20</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload></answerpreload>
    <useace>1</useace>
    <resultcolumns></resultcolumns>
    <template><![CDATA[import re
import subprocess
import sys

# --- checks begin ---
# Shared verbatim with prototype-go-testwriter.xml and executed by
# tools/simulate-grader-go.py, which rejects any drift between the copies.
TOKENS = re.compile(
    r'`[^`]*`|"(?:\\.|[^"\\\n])*"|\'(?:\\.|[^\'\\\n])*\'|//[^\n]*|/\*.*?\*/', re.S)


def _strip_all(src):
    return TOKENS.sub(' ', src)


def _strip_comments(src):
    return TOKENS.sub(lambda m: m.group(0) if m.group(0)[0] in '"\'`' else ' ', src)


OPERATORS = {
    '==': r'(?<![=!<>])==(?!=)',
    '!=': r'!=',
    '+': r'(?<!\+)\+(?!\+)',
    '&': r'(?<!&)&(?!&)',
    '&&': r'&&',
    '||': r'\|\|',
    '!': r'!(?!=)',
    '<-': r'<-',
    ':=': r':=',
    '%': r'(?<!%)%(?!=)',
    '*': r'(?<![*/])\*(?![*/=])',
    'nakedreturn': r'\breturn[ \t]*(?:\n|\})',
}


def _pattern(tok):
    if tok in OPERATORS:
        return OPERATORS[tok], 'stripped'
    if re.fullmatch(r'\d+', tok):
        return r'(?<![\w.])' + tok + r'(?![\w.])', 'comments'
    if tok in ('"', "'", '`') or tok.startswith('%'):
        return re.escape(tok), 'comments'
    return r'\b' + re.escape(tok) + r'\b', 'stripped'


def _max_body_lines(src):
    worst = 0
    for m in re.finditer(r'\bfunc\b[^{]*\{', src):
        i, depth = m.end(), 1
        while i < len(src) and depth:
            if src[i] == '{':
                depth += 1
            elif src[i] == '}':
                depth -= 1
            i += 1
        body = src[m.end():i - 1]
        worst = max(worst, len([l for l in body.split('\n') if l.strip()]))
    return worst


def check_rules(student, banned, required, limit):
    """Return None when every rule holds, else the failure message."""
    views = {'stripped': _strip_all(student), 'comments': _strip_comments(student)}
    for tok in banned:
        pat, view = _pattern(tok)
        if re.search(pat, views[view]):
            return 'banned token present: ' + tok
    for tok in required:
        pat, view = _pattern(tok)
        if not re.search(pat, views[view]):
            return 'required token missing: ' + tok
    for item in limit:
        kind, _, n = item.partition('=')
        if kind == 'return':
            count = len(re.findall(r'\breturn\b', views['stripped']))
        elif kind == 'bodylines':
            count = _max_body_lines(views['stripped'])
        else:
            return 'unknown limit: ' + item
        if count > int(n):
            return 'limit exceeded: ' + item
    return None
# --- checks end ---

student = """{{ STUDENT_ANSWER | e('py') }}"""
banned = """{{ QUESTION.parameters.banned | default('') | e('py') }}""".split()
required = """{{ QUESTION.parameters.required | default('') | e('py') }}""".split()
limit = """{{ QUESTION.parameters.limit | default('') | e('py') }}""".split()

SEP = "#<ab@17943918#@>#"
testcodes = []
stdins = []
{% for TEST in TESTCASES %}
testcodes.append("""{{ TEST.testcode | e('py') }}""")
stdins.append("""{{ TEST.stdin | e('py') }}""")
{% endfor %}

failure = check_rules(student, banned, required, limit)
if failure is not None:
    print(failure)
    for _ in testcodes[1:]:
        print(SEP)
    sys.exit(0)

harness = ['package main\n\nimport (\n\t"fmt"\n\t"os"\n)\n\nvar _ = fmt.Sprint\n\nfunc main() {\n\tswitch os.Args[1] {\n']
for i, code in enumerate(testcodes):
    harness.append('\tcase "%d":\n' % i)
    harness.append(code)
    harness.append('\n')
harness.append('\t}\n}\n')

with open('student.go', 'w') as f:
    f.write(student)
with open('tests.go', 'w') as f:
    f.write(''.join(harness))

env = {
    'GOCACHE': '/tmp/.go-cache',
    'GOMAXPROCS': '1',
    'PATH': '/usr/bin:/bin',
    'HOME': '/tmp',
}
cp = subprocess.run(['go', 'build', '-p=1', '-o', 'prog', 'student.go', 'tests.go'],
                    env=env, capture_output=True, text=True)
if cp.returncode != 0:
    print('** Compilation failed **')
    print(cp.stderr)
    for _ in testcodes[1:]:
        print(SEP)
    sys.exit(0)

for i, (code, stdin_text) in enumerate(zip(testcodes, stdins)):
    if code.strip() == '// rules':
        print('rules ok')
    else:
        r = subprocess.run(['./prog', str(i)], input=stdin_text, env=env,
                           capture_output=True, text=True)
        sys.stdout.write(r.stdout)
        if r.returncode != 0:
            sys.stdout.write(r.stderr)
    if i != len(testcodes) - 1:
        print(SEP)
]]></template>
    <iscombinatortemplate>1</iscombinatortemplate>
    <allowmultiplestdins>1</allowmultiplestdins>
    <answer></answer>
    <validateonsave>0</validateonsave>
    <testsplitterre><![CDATA[|#<ab@17943918#@>#\n|ms]]></testsplitterre>
    <language>python3</language>
    <acelang>go</acelang>
    <sandbox></sandbox>
    <grader>EqualityGrader</grader>
    <cputimelimitsecs>20</cputimelimitsecs>
    <memlimitmb>16000</memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <testcases>
    </testcases>
</question>
</quiz>
